Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`.

* You will be filling the parts denoted with "...". You do not need to change anything else in the code.
* If you see a "All tests passed!" message, then you are good to go. If you receive an error, you might need to check your code again.
* We are also checking your code with hidden cases! Unless specifically stated, always use the variables given in the function definitions, and **do not** return directly calculated values. These will definitely fail the test cases.

---

```python
import math
```

---
## Question 1

An encryption-decryption system consists of three elements: encode, transmit, and decode. A faulty encode occurs in 0.8% of the messages processed, transmission errors occur in 1% of the messages, and a decode error occurs in 0.1% of the messages. Assume
the errors are independent. Round your answers to four decimal places (e.g. 98.7654).

### Part A (20 points)

What is the probability of a completely defect-free message?

<!-- grade_id: cell-5f2518c7755cd892, points: 0, grade: false, solution: true -->
```python
def q1_A(E, T, D):
    '''
    Inputs:
    E: Percentage of faulty encode 
    T: Percentage of faulty transmit
    D: Percentage of faulty decode
    Output:
    prob: Probability of defect-free message
    '''
    # YOUR CODE HERE
    prob = ...
    
    return prob
```

<!-- grade_id: cell-d168a51e59343992, points: 20, grade: true, solution: false -->
```python
# DON'T CHANGE THIS CELL!

try:
    assert math.isclose(q1_A(0.008, 0.01, 0.001), 0.9811, abs_tol=1e-3)
    assert math.isclose(q1_A(0.1, 0.02, 0.005), 0.8776, abs_tol=1e-3)
    

    print("All tests passed!")

except AssertionError as e:
    print("There are failing test cases!")
```

### Part B (20 points)

What is the probability of a message that has either an encode or a decode error?

<!-- grade_id: cell-374954e369b8e5c8, points: 0, grade: false, solution: true -->
```python
def q1_B(E, T, D):
    '''
    Inputs:
    E: Percentage of faulty encode 
    T: Percentage of faulty transmit
    D: Percentage of faulty decode
    Output:
    prob: Probability of a message that has either and encode or a decode error
    '''
    # YOUR CODE HERE
    prob = ...
    
    return prob
```

<!-- grade_id: cell-6b43f920ccb3f6ab, points: 20, grade: true, solution: false -->
```python
# DON'T CHANGE THIS CELL!

try:
    assert math.isclose(q1_B(0.008, 0.01, 0.001), 0.009, abs_tol=1e-3)
    assert math.isclose(q1_B(0.1, 0.02, 0.005), 0.1045, abs_tol=1e-3)


except AssertionError as e:
    print("There are failing test cases!")
```

---

## Question 2

Barcelona (Barca) football team has won 78% of the games in its recent history. In the games in which Messi played, Barca has won the game in 94% of times. It is also known that Messi has played in 82% of Barca’s games. Using the above information, answer the following question:

### Part A (15 points)

What is the probability of Messi played in the game and Barca lost?

<!-- grade_id: cell-71cd3f1d848ddc13, points: 0, grade: false, solution: true -->
```python
def q2_A(W, M, WM):
    '''
    Inputs:
    W: Percentage of games Barcelona won
    M: Percentage of games Messi played
    WM: Percentage of games Barcelona won in which Messi played
    Output:
    prob: Probability of Messi played in the game and Barca lost
    '''
    # YOUR CODE HERE
    prob = ...
    
    return prob
```

<!-- grade_id: cell-21829b2ccae23071, points: 15, grade: true, solution: false -->
```python
# DON'T CHANGE THIS CELL!

try:
    assert math.isclose(q2_A(0.78, 0.82, 0.94), 0.0492, abs_tol=1e-3)
    assert math.isclose(q2_A(0.1, 0.02, 0.005), 0.0199, abs_tol=1e-3)
    

except AssertionError as e:
    print("There are failing test cases!")
```

### Part B (15 points)

What is the probability of Messi didn’t play in the game and Barca won?

<!-- grade_id: cell-976296c4d694825b, points: 0, grade: false, solution: true -->
```python
def q2_B(W, M, WM):
    '''
    Inputs:
    W: Percentage of games Barcelona won
    M: Percentage of games Messi played
    WM: Percentage of games Barcelona won in which Messi played
    Output:
    prob: Probability of Messi didn’t play in the game and Barca won
    '''
    # YOUR CODE HERE
    W_and_M = ...
    prob = ...
    
    return prob
```

<!-- grade_id: cell-2f17b7d16feea3db, points: 15, grade: true, solution: false -->
```python
# DON'T CHANGE THIS CELL!

try:
    assert math.isclose(q2_B(0.78, 0.82, 0.94), 0.0092, abs_tol=1e-3)
    assert math.isclose(q2_B(0.1, 0.02, 0.005), 0.0999, abs_tol=1e-3)
except AssertionError as e:
    print("There are failing test cases!")
```

### Part C (15 points - Manual Grading)

Draw a probability tree by using the given information.

! Upload your drawings to Blackboard as a seperate file.

### Part D (15 points)

You missed Barca’s evening game, but you have just read in the morning paper that Barca could not win the game. What is the probability that Messi did not play in that game?

! **Only for this question**, use the values you calculated directly; instead of using variables. Use the values from your probabiltiy tree from Part C.

<!-- grade_id: cell-5878781fdaefbf94, points: 0, grade: false, solution: true -->
```python
def q2_D(W, M, WM):
    '''
    Inputs:
    W: Percentage of games Barcelona won
    M: Percentage of games Messi played
    WM: Percentage of games Barcelona won in which Messi played
    Output:
    prob: Probability that Messi did not play in that game
    '''
    # YOUR CODE HERE
    prob = ...
    
    return prob
```

<!-- grade_id: cell-82fed84e81bf7c0b, points: 15, grade: true, solution: false -->
```python
# DON'T CHANGE THIS CELL!

try:
    print(f"Your answer is: {q2_D(0.78, 0.82, 0.94):.4f}.\nThere is no visible tests for this question.")

except AssertionError as e:
    print("There are failing test cases!")
```